```python
# Initialize Otter
import otter
grader = otter.Notebook("demo.ipynb")
```

# Otter-Grader Tutorial

This notebook is part of the Otter-Grader tutorial. For more information about Otter, see our [documentation](https://otter-grader.rtfd.io).

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Write a function `square` that returns the square of its argument.

```python
def square(x):
    ...
```

```python
grader.check("q1")
```

**Question 2:** Write an infinite generator of the Fibonacci sequence `fiberator` that is *not* recursive.

```python
def fiberator():
    ...
```

```python
grader.check("q2")
```

**Question 3:** Create a DataFrame mirroring the table below and assign this to `data`. Then group by the `flavor` column and find the mean price for each flavor; assign this **series** to `price_by_flavor`.

| flavor | scoops | price |
|-----|-----|-----|
| chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| chocolate | 3 | 5 |

```python
data = ...
price_by_flavor = ...
price_by_flavor
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4:** Create a barplot of `price_by_flavor`.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5:** What do you notice about the bar plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: None)
```
>>> square(1) == 1
True
```
```
>>> square(0) == 0
True
```


Tests for q2 (points: None)
```
>>> f = fiberator()
>>> assert next(f) == 0
>>> assert next(f) == 1
```


Tests for q3 (points: None)
```
>>> len(data["flavor"].unique()) == 4
True
```
```
>>> for l in ["chocolate", "vanilla", "strawberry", "mint"]:
...     assert l in data["flavor"].unique()
```
```
>>> assert type(price_by_flavor) == pd.Series
```
```
>>> assert len(price_by_flavor) == 4
```
